<a href="https://colab.research.google.com/github/vsaathwika0904/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE=json.load(open("allergies.json"))
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_count(file):
  """
  (dict)-> int
  We confirmed the allergies by the status of the patient which is in the resource to
  help us in counting the allergies.The resource data is in a dictionary, in order
  to call all entries we are using for loop to call them. In the allergies file
  when the status shows that it is confirmed then we are incrementing the count
  by +1 and then we'll return the count.

  >>> allergy_count(ALLERGIES_FILE)
  4

  """
  count=0
  for entries in range(len(ALLERGIES_FILE['entry'])):
    if ALLERGIES_FILE['entry'][entries]['resource']['status']=='confirmed':
      count+=1
  return count
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    allergy_count(ALLERGIES_FILE)
Expecting:
    4
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here
def patient_count(file):
  """
  (dict) -> int
  Here to find the number of patients, we are first creating an empty list to store the patient
  names then we will be using for loop to call all the entries. Then, with the empty list we
  created before, we will append it with the allergies file and for not to repeat the names, we
  have to create a set for the patient list. In the end we will return the length of the patient list.

  >>> patient_count(ALLERGIES_FILE)
  2

  """
  patient_list=[]
  for entries in range(len(ALLERGIES_FILE['entry'])):
   patient_list.append(ALLERGIES_FILE['entry'][entries]['resource']['patient']['display'])
  patient_list=set(patient_list)
  return len(patient_list)
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

```python
import doctest
doctest.run_docstring_examples(patient_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count(ALLERGIES_FILE)
Expecting:
    2
ok
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
def allergy_per_patient(file):
  """
  (dict) -> dict
  Here they are asking how many allergies each patient has so first we will create an empty
  list called patient list and next we will create a dictionary to note the allergies per
  patient. To call all the entries we are going to create an for loop and we are appending
  the empty patient list with the patient names in each entry and for not to repeat the names, we have
  to create a set for the patient list to provide unique names. Next we will be creating
  two for loops one for the unique name list and other for the patient list. If the name
  is equal to the names in patient list then we will increment the count with +1 and for
  the dictionary previously created we will represent it by no.of allergies of the patient
  and then we will return the allergy_counts.

  >>> allergy_per_patient(ALLERGIES_FILE)
  {'Paul Boal': 1, 'Jason Argonaut': 3}

  """
  patient_list=[]
  allergy_counts={}
  for entries in range(len(ALLERGIES_FILE['entry'])):
   patient_list.append(ALLERGIES_FILE['entry'][entries]['resource']['patient']['display'])
  uniquename_list=set(patient_list)
  for names in uniquename_list:
    count=0
    for i in patient_list:
      if names==i:
        count+=1
    allergy_counts[names]=count
  return allergy_counts

```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Paul Boal': 1, 'Jason Argonaut': 3}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

```python
import doctest
doctest.run_docstring_examples(allergy_per_patient, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_per_patient(ALLERGIES_FILE)
Expecting:
    {'Paul Boal': 1, 'Jason Argonaut': 3}
ok
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(file):
  """
  (dict) -> list
  First we are creating an empty list to store the patient name, what they are allergy to and
  the reaction to their allergies. Now we are calling all the entries using a for loop. As we
  have to provide the patient name, what they are allergy to and the reaction to their allergies
  so then we are finding the location of all these data and we are extracting them. Next we need
  to append the created empty list with the extracted data and we have to return the created
  patient list
  >>> allergy_list(ALLERGIES_FILE)
  [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]

  """
  patient_list=[]
  for entries in range(len(ALLERGIES_FILE['entry'])):
    patient=ALLERGIES_FILE['entry'][entries]['resource']['patient']['display']
    substance=ALLERGIES_FILE['entry'][entries]['resource']['substance']['text']
    reaction=ALLERGIES_FILE['entry'][entries]['resource']['reaction'][0]['manifestation'][0]['text']
    patient_list.append([patient,substance,reaction])
  return patient_list
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

```python
import doctest
doctest.run_docstring_examples(allergy_list, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_list(ALLERGIES_FILE)
Expecting:
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]
ok
```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(file, patient_name='name',drug='medicine',):
  """
(dict) -> str
 First we are creating a variable report as none and then to call all entries we are using a for loop.
 Next we are finding the locations of patient and substance data to extract them. Now if the name of
 the patient name and drug is equal to name we get while the loop is running then we will return it
 to report which we extracted from manifestation text, if the patient name and drug does not match
 with the data in our file then the loop will stop and it will return as none.

 >>> allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'

 """
  report=None
  for entries in range(len(ALLERGIES_FILE['entry'])):
    patient=ALLERGIES_FILE['entry'][entries]['resource']['patient']['display']
    substance=ALLERGIES_FILE['entry'][entries]['resource']['substance']['text']
    if patient==patient_name and drug==substance:
       report=ALLERGIES_FILE['entry'][entries]['resource']['reaction'][0]['manifestation'][0]['text']
       return report
  if report==None:
    return report
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python
import doctest
doctest.run_docstring_examples(allergy_reaction, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
Expecting nothing
**********************************************************************
File "__main__", line 11, in NoName
Failed example:
    allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
Expected nothing
Got:
    True
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

